# <center> ЛАБОРАТОРНАЯ  РАБОТА 2 </center>
# <center> РЕШЕНИЕ  НЕЛИНЕЙНЫХ  УРАВНЕНИЙ </center>

Теоретический материал к данной теме содержится в [1, глава 4].

```python
from sympy.solvers import solve
from sympy import exp, sin, cos, sqrt, log, ln, pi
from sympy import Rational as syR
from sympy import Symbol, diff
from scipy.misc import derivative

import numpy as np

import matplotlib.pyplot as plt
%matplotlib inline
```

## Задание 1
Даны два уравнения  f(x)=0 и g(x)=0. Найти с точностью $10^{-10}$  все корни уравнений, содержащиеся на отрезке [a, b]. Для решения  задачи использовать  метод бисекции.  Найти корни с помощью встроенной функции.
## ПОРЯДОК  РЕШЕНИЯ  ЗАДАЧИ:
1. Найти аналитическое решение уравнения  f(x)=0.
2. Локализовать корни  f(x)=0   графически.
3. Найти корни уравнения f(x)=0  с точностью    с помощью метода бисекции.
4. Используя  встроенную функцию найти корни уравнения f(x)=0 с точностью  .
5. Аналогично п. 1-4 попытаться найти корни уравнения  g(x)=0. Объяснить полученные результаты.

Ваши функции:  
$f(x) = {Task1[TASK_VARIANT][f_latex]}$  
$g(x) = {Task1[TASK_VARIANT][f_latex]}$

Пределы:  
$[a;b] = [ {Task1[TASK_VARIANT][a]};{Task1[TASK_VARIANT][b]} ] $

```python
epselon1 = 10**(-10)
# Пределы:  
a1,b1 = {Task1[TASK_VARIANT][a]}, {Task1[TASK_VARIANT][b]}
```

Запишите $f(x)$ и $g(x)$ как python'ские функции

<!-- grade_id: cell-7d3d5f24ad1b0a99, points: 1, grade: true, solution: true -->
```python
def f1(x):
    ### BEGIN SOLUTION
    return x
    ### END SOLUTION
```

Output:
```
File "<ipython-input-1-2a8e5b0b810a>", line 1
    def f(x):
             ^
SyntaxError: unexpected EOF while parsing
```

<!-- grade_id: cell-8238d2875eade012, points: 1, grade: true, solution: true -->
```python
def g1(x):
    ### BEGIN SOLUTION
    return x
    ### END SOLUTION
```

```python
Локализуйте корни уравнения графически, используя пакет matplotlib
```

<!-- grade_id: cell-40a778165eca267f, points: 1, grade: true, solution: true -->
```python
## ЛОКАЛИЗУЙТЕ КОРНИ ГРАФИЧЕСКИ ##
```

Допишите функцию BisecMethod, которая реализует метод бисекции.  
BisecMethod - должен вохвращать точку и количество шагов итераций, необходимое для достижения заданной точности

<!-- grade_id: cell-c4208357eb7a6a8c, points: 1, grade: true, solution: true -->
```python
def BisecMethod():
    ### BEGIN SOLUTION
    return x
    ### END SOLUTION
```

```python
## Тут решение ##
bisec_method_x_1, bisec_method_n_1  = BisecMethod(f1,a1,b1,epselon1)
```

<!-- grade_id: cell-8f289dd75e7d7d43, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(abs(f1(bisec_method_x_1)) < epselon1, True)

### BEGIN HIDDEN TESTS
#assert_equal(square(5),[25])
### END HIDDEN TESTS
```

```python
## Решение встроенной функцией ##
integrated_x_1 = 
```

```python
## Тест ##
from numpy.testing import assert_allclose


assert_allclose(abs(f1(integrated_x_1)), 0, atol=epselon1)
```

```python
## Простор для функции g ##
```

```python
## Простор для функции g ##
```

```python
## Простор для функции g ##
```

<!-- grade_id: cell-676e1782fafd4367, points: 1, grade: true, solution: true -->
Что не так с функцией $g(x)$? 

Запишите свой ответ здесь

## Задание 2  
Найти указанный в варианте  корень уравнения $f(x)=0$ с точностью $10^{-6}$ ,  двумя  способами.
   
   а) Использовать метод бисекции. Предварительно определить отрезок локализации  [a, b].
   
   b) Использовать метод Ньютона. В качестве начального приближения для метода Ньютона взять  середину отрезка локализации из п. а).
   
   Сравнить число итераций в п. a), b).


Ваша функция $f(x) = {Task2[TASK_VARIANT][f_latex]}$  
Корень, который Вам нужен: {Task2[TASK_VARIANT][which]}

```python
epselon2 = 10**(-6) 
```

Запишите $f(x)$ как python'ские функцию

<!-- grade_id: cell-9b3935eae80a2832, points: 1, grade: true, solution: true -->
```python
def f2(x):
    ### BEGIN SOLUTION
    return x
    ### END SOLUTION
```

Определите отрезок локализации  
Сохраните как $a2,b2$

<!-- grade_id: cell-11b96dfd585527d9, points: 1, grade: true, solution: true -->
```python
## Определите отрезок локализации ##
```

Решите задачу используя метод бисекции

```python
bisec_method_x_2, bisec_method_n_2  = BisecMethod(f2,a2,b2,epselon2)
```

<!-- grade_id: cell-ec87b93c2f0ab8c2, points: 1, grade: true, solution: false -->
```python
assert_equal(abs(f2(bisec_method_x_2)) < epselon2, True)
```

Определите метод Ньютона решения нелинейных уравнений  
Он должен возвращать два числа - точку x и количество итераций

<!-- grade_id: cell-ed02de9cea0bc431, points: 1, grade: true, solution: true -->
```python
def NewtoonMethod():
    ### BEGIN SOLUTION
    return x
    ### END SOLUTION
```

<!-- grade_id: cell-3840b5bf792e66e9, points: 1, grade: true, solution: false -->
```python
newtoon_method_x_2, newtoon_method_n_2  = NewtoonMethod()
## Тут тест ##
assert_equal(abs(f2(newtoon_method_x_2)) < epselon2, True)
```

Сравните число итераций в двух методах

```python
## ТЕСТ ДЛЯ ЧИСЛА ИТЕРАЦИЙ ##
assert_equal(anewtoon_method_n_2 < bisec_method_n_2, True)
```

## Задание 3

Локализовать корни уравнения $f(x)=0$ и найти их с точностью $10^{-5}$ , используя метод простой итерации. К виду $x=\phi(x)$, удобному для итераций, уравнение $ f(x)=0$ привести двумя способами.

a) Преобразовать уравнение к виду  $x=x-\alpha f(x)$,    где $\alpha=\frac{2}{M+m}$, $0 < m \leq f'(x) \leq M$, а $x$ принадлежит отрезку    локализации [a, b].
     
b) Любым  другим  преобразованием  уравнения. Проверить достаточное условие сходимости метода. 
     
Использовать критерий окончания итерационного процесса вида  , где  в  п. a)  $q=\frac{(M-m)}{(M+m)}$, в п. b) $q = \max\limits_{x \in [a,b]}{\phi'(x)}$


Ваша функция:  
$ f(x) = {Task3[TASK_VARIANT][f_latex]} $ 

```python
epselon3 = 10**(-5)
```

Запишите $f(x)$ как python'ские функцию

<!-- grade_id: cell-3fc108b2ecafe87a, points: 1, grade: true, solution: true -->
```python
def f3(x):
    ### BEGIN SOLUTION
    return x
    ### END SOLUTION
```

Реализуйте метод простой итерации

<!-- grade_id: cell-82683c24bbaeb041, points: 1, grade: true, solution: true -->
```python
def SimpleIteration():
    ### BEGIN SOLUTION
    return x
    ### END SOLUTION
```

<!-- grade_id: cell-25098480fe60ae90, points: 0, grade: false, solution: true -->
```python
## Тут все для SimpleIteration
simple_iteration_x_3_1 = SimpleIteration()
```

<!-- grade_id: cell-6b6542dfa48228be, points: 1, grade: true, solution: false -->
```python
## Первое решение тест ##
assert_equal(abs(f3(simple_iteration_x_3_1)) < epselon3, True)
```

<!-- grade_id: cell-8b14300694c8b4d3, points: 0, grade: false, solution: true -->
```python
## Тут все для SimpleIteration
simple_iteration_x_3_2 = SimpleIteration()
```

<!-- grade_id: cell-b3a075a71af17143, points: 1, grade: true, solution: false -->
```python
## Второе решение тест ##
assert_equal(abs(f3(simple_iteration_x_3_2)) < epselon3, True)
```

Проверьте условие сходимости метода.

<!-- grade_id: cell-7882ce3ca2938ab6, points: 0, grade: true, solution: true -->
```python

```

#### ЛИТЕРАТУРА:
1. Амосов А.А., Дубинский Ю.А., Копченова Н.В. Вычислительные методы для инженеров. М.: Высшая школа, 1994.
